---
title: 'Systems of Linear Equations'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.

`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```

---

## Least Squares Approximation

```yaml
type: VideoExercise
key: b4578b9151
xp: 50
```

`@projector_key`
94b2bb832f5e9fe7d78c69b3c169956e
